# Morse Code #

This notebook covers some basic functions of Python, and use them to encrypt and decrypt Morse code.

### History of Morse Code ###

Invented by Samuel Morse in 1830s and 1840s,telecommunication took a large leap in long-distance communication. Morse code assigns a set of dots and dashes to each letter in the alphabet. The first telegraph message was sent from Washington D.C. to Baltimore, Maryland in 1844. The Morse Code was eventually replaced by telephones in the 21st Century. 
The alphabet and the associated Morse code is shown below: <br>
<img src="https://miro.medium.com/max/700/1*1mOO6Lmh9ASEp_BRZvOJKQ.jpeg"
     alt="Markdown Monster icon"
     style="float: center; margin-right: 0px;"
     width = '400px'
     height = '300px'
     />

### Dicitionary ###

When creating a list of variables, coders will often group related variables (such as the letters in the alphabet) into a **dictionary**. A dictionary has words (called keys), with definitions (called values). Dictionaries in Python are defined using "{'key1':'value1', 'key2': 'value2'}" Check out the example below:

<pre><code>dictionary = {'key1': 'value1'.
              'key2': 'value2'
             }</code></pre>

### Let's try an example for creating the Morse code alphabet ###
Let's try to work out an example. Write a dictionary in the cell below for the Morse code alphabet. 
Set the keys to be the English alphabet, and set the values to be the Morse code translation. 
(Example: 'A': '.-')

```python
# Dictionary for Morse Code #

### BEGIN CODE HERE ###

### END CODE HERE ###
```

### Conditional Statements - "If/Elif/Else" ###
One of the most useful coding functions are if-elif-else statements.  Each code block starts with a conditional statement, ending with ":". The proceeding "then" statements are indented. See the below example for the conditional statement format:
<pre><code>if statement_1 is true:
    action1
    action2
elif statement_2 is true:
    action3
else:
    action 4</code></pre>
The "If" statement is read first. If the first conditional statement is true, the indented code (i.e. action1 and action2) are carried out. The other actions under the other two statements are not carried out. If the first conditional statement isn't true, then the second conditional statement is read, and if true, the indented code (action3) is carried out intsead. If none of the if & elif statements are true, then the code underneath the else statement is carried out.




### Let's try an example to create conditional statements ###
Write an if-else statement where it will print out "true" when the "if" statement is true, and will print out "false" when "else" is true

```python
# If-Else Example #

### BEGIN CODE HERE ###

### END CODE HERE ###
```

### For Loop ### 
Another useful function in Python is the "for" loop. In a "for" loop, code is repeated for a set amount of iterations, as defined in the "for" statement, which includes ":" at the end.

<pre><code>for x in variable: --> This line means for each number of variables in "txt" variable:
    perform a function </pre></code>
    
Sample code:
<pre><code>for x in input:
    print('Hello')</pre></code>
    
The above code will print out the code for however many variables are in text. 


NOTE: a useful feature in "for" loops. If you want to add multiple characters to a single variable, you can do the following:

<pre><code>output = ''
dictionary = {'key1': 'value1','key2':'value2'}
for x in input:
    output += dictionary.get(x)</pre></code>

The above code will extract the appropriate value from the dictionary based on the input key "x" in the for loop

#### Try coding a for loop that prints out the numbers 1-10 , one at a time.

Remember to use the syntax above. You can loop through a list variable (that is, a variable with a list or list like data type, such as an array or a string, which is a "list" of characters.), or you can use "in range(x)", where x is the number of times you wish to run the loop. Which do you believe is more suitable for this task?

```python
def loop():
    
    ##Your code here.##
    
    return 
```

### Useful functions for Morse code ###

* **startswith(prefix[,start[,end]])**: checks if a string "prefix" if beginning position (start) and ending position (end)
* **split(' ')**: this function will separate by a separator varible. If it is empty quotation marks, then the variable will be separated by empty space.
* **get( )**: this function retrives the value of a defined key. For example, if a variable var = 'A', A.get() will retrive the Morse code value associated with A. 
* **strip( )**: this function removes any spaces at the beginning or at the end of the string. 

To switch the keys/values in a dictionary value:
<pre><code></code> dict([(v, k) for k, v in d.items()]) </code></pre> where v is value, k is key

Now we are going to develop a Morse code function that will encrypt and decrypt Morse code. You may use the examples used in the previous sections as part of your funciton. You will be encrypting and decrypting the examples below:
<pre><code>print(morse('python'))
# .--. -.-- - .... --- -.
print(morse('.--. -.-- - .... --- -.'))
# PYTHON
print(morse(morse('HEY')))
# HEY</code></pre>


<!-- grade_id: c8, points: 0, grade: false, solution: true -->
```python
# Creating function for Morse Code 
def morse(txt):
    '''Morse code encryption and decryption'''
    
    ### BEGIN DICTIONARY DEFINITION CODE ###
    # Write a dictionary definition for the Morse code #

    ### END DICTIONARY DEFINITION CODE ### 
    
    translation = '' # Let this be the output variable for this function morse()
    
    ### BEGIN ENCRYPTION CODE ###
    # Write an if statement that detects if the input variable starts with a "." or "-". 
    # Write a variable will switch the key and values of the dictionary defition defined above #
    # Write a line of code that will separate each character in the input variable by a empty space
    # For each character, get the appropriate value from the new encryption dictionary. Add this value to output#
    
    ### END ENCRYPTION CODE ### 
    
    ### BEGIN DEENCRYPTION CODE ### 
    # Write an else statement# 
    # Add a line that converts the input variable to upper case letters
    # For each character, get the appropriate value from the decryption dictionary defined above.Concentate an empty space at the end of each dot and dash. # 
    ### END DEENCRYPTION CODE ### 
    
    ### BEGIN OUTPUT CODE ###
    # Write a line of code that returns the translation for encryption or decryption with the empty spaces removed. 
    ### END OUTPUT CODE ###
```

```python
def morse(txt):
    '''Morse code encryption and decryption'''
    
    d = {'A':'.-','B':'-...','C':'-.-.','D':'-..','E':'.',
         'F':'..-.','G':'--.','H':'....','I':'..','J':'.---',
         'K':'-.-','L':'.-..','M':'--','N':'-.','O':'---',
         'P':'.--.','Q':'--.-','R':'.-.','S':'...','T':'-',
         'U':'..-','V':'...-','W':'.--','X':'-..-','Y':'-.--',
         'Z':'--..', ' ':'.....'}
    translation = ''
    
    # Encrypt Morsecode
    if txt.startswith('.') or txt.startswith('−'):
        # Swap key/values in d:
        d_encrypt = dict([(v, k) for k, v in d.items()])
        # Morse code is separated by empty space chars
        txt = txt.split(' ')
        for x in txt:
            translation += d_encrypt.get(x)
        
    # Decrypt to Morsecode:
    else:
        txt = txt.upper()
        for x in txt:
            translation += d.get(x) + ' '
    return translation.strip()
print(morse('python'))
# .--. -.-- - .... --- -.
print(morse('.--. -.-- - .... --- -.'))
# PYTHON
print(morse(morse('HEY')))
# HEY

```

Output:
```
.--. -.-- - .... --- -.
PYTHON
HEY
```

### TEST YOUR CODE ###
Now let's test your code below. Run the cells below to check if they match. If there are no errors, then you successfully wrote a Morse code encryption and decryption function. 

<!-- grade_id: c11, points: 0, grade: true, solution: false -->
```python
from nose.tools import assert_equal

assert_equal(morse('python'),'.--. -.-- - .... --- -.')
assert_equal(morse('.--. -.-- - .... --- -.'),'PYTHON')
assert_equal(morse(morse('HEY')),'HEY')
```

### Congratulations ! You have successfully completed the Morse Code Section ###