### Instructions ###
Answer each of the questions by entering the answer value into the "input_var" variable.  <br>
You can ignore the "Test your code from above here" cells as they will be used to grade each question. 

### Question 1: What is PYNQ? ###
<input type="checkbox" name="choice" value="1"> A development board <br> 
<input type="checkbox" name="choice" value="1"> An operating system <br> 
<input type="checkbox" name="choice" value="1"> An open-source project and framework <br> 
<input type="checkbox" name="choice" value="1"> A Python-based game

```python
# Enter 1, 2, 3, or 4 as the value of the variable below: i.e. input_var = 1
input_var = 1
```

<!-- grade_id: q1, points: 3, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def input_value():
    return input_var

from nose.tools import assert_equal
assert_equal(input_value(), 3)

### END HIDDEN TESTS
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
<ipython-input-2-58a3b557131c> in <module>()
      3     return input_var
      4 
----> 5 from nose.tools import assert_equal
      6 assert_equal(input_value(), 3)
      7 

ModuleNotFoundError: No module named 'nose'
```

### Question 2: Which of the following environments is used with PYNQ? ###
<label for="option 1">1. PyCharm</label><br>
<label for="option 2">2. Jupyter Notebooks</label><br>
<label for="option 3">3. Visual Studio</label><br>
<label for="option 4">4. BlueJ</label>

Enter 1, 2, 3, or 4 in the cell below

```python
# Enter 1, 2, 3, or 4 as the value of the variable below: i.e. input_var = 1
input_var = 
```

<!-- grade_id: q2, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def input_value():
    return input_var

from nose.tools import assert_equal
assert_equal(input_value(), 1)
assert_equal(input_val)

### END HIDDEN TESTS
```

### Question 3: What is the standard way to create a PYNQ image for a custom board? ###
<label for="option 1">1. Via Petalinux HDF file</label><br>
<label for="option 2">2. Via Petalinux BOOT.BIN file</label><br>
<label for="option 3">3. Via Petalinux BSP file</label><br>
<label for="option 4">4. Via Petalinux image.ub</label>

```python
# Enter 1, 2, 3, or 4 as the value of the variable below: i.e. input_var = 1
input_var = 
```

<!-- grade_id: q3, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def input_value():
    return input_var

from nose.tools import assert_equal
assert_equal(input_value(), 2)

### END HIDDEN TESTS
```

### Question 5: With PYNQ, hardware system designs are referred to as: ###
<label for="option 1">1. APIs</label><br>
<label for="option 2">2. Layers</label><br>
<label for="option 3">3. Notebooks</label><br>
<label for="option 4">4. Overlays</label>

```python
# Enter 1, 2, 3, or 4 as the value of the variable below: i.e. input_var = 1
input_var = 
```

<!-- grade_id: q5, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def input_value():
    return input_var

from nose.tools import assert_equal
assert_equal(input_value(), 4)

### END HIDDEN TESTS
```

### Question 6: What set of files are necessary to deploy a PYNQ system? ###
<label for="option 1">1. Bitstream (*.bit), Hardware Handoff (*.hwh), Jupyter Notebook (*.ipynb)</label><br>
<label for="option 2">2. Bitstream (*.bit), Hardware Handoff (*.hwh), Solution (.sln)</label><br>
<label for="option 3">3. Hardware Handoff (*.hwh), Solution (*.sln), Jupyter Notebook (*ipynb)</label><br>
<label for="option 4">4. Hardware Handoff (*hwh), Java Archive (*.jar), Text (.txt)</label>

```python
# Enter 1, 2, 3, or 4 as the value of the variable below: i.e. input_var = 1
input_var = 
```

<!-- grade_id: q6, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def input_value():
    return input_var

from nose.tools import assert_equal
assert_equal(input_value(), 1)

### END HIDDEN TESTS
```

### Question 7: Which command is used to install new PYNQ projects using hybrid packages? ###
<label for="option 1">1. Conda install</label><br>
<label for="option 2">2. PIP install</label><br>
<label for="option 3">3. Apt-get install</label><br>
<label for="option 4">4. Easy-install</label>

```python
# Enter 1, 2, 3, or 4 as the value of the variable below: i.e. input_var = 1
input_var = 
```

<!-- grade_id: q7, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
def input_value():
    return input_var

from nose.tools import assert_equal
assert_equal(input_value(), 2)

### END HIDDEN TESTS
```

Output:
```
/usr/lib/python3.8/unittest/case.py in assertEqual(self, first, second, msg)
    910         """
    911         assertion_func = self._getAssertEqualityFunc(first, second)
--> 912         assertion_func(first, second, msg=msg)
    913 
    914     def assertNotEqual(self, first, second, msg=None):

/usr/lib/python3.8/unittest/case.py in _baseAssertEqual(self, first, second, msg)
    903             standardMsg = '%s != %s' % _common_shorten_repr(first, second)
    904             msg = self._formatMessage(msg, standardMsg)
--> 905             raise self.failureException(msg)
    906 
    907     def assertEqual(self, first, second, msg=None):

AssertionError: 1 != 2
```